# Week1 Assignment - Testing OKPy and the autograder.

For this assignment, you will load the iris flowers dataset from sklearn and predict flower species based on their petal and sepal sizes.

```python
"""This is a docstring!"""
```

$\beta=5$

```python
alpha = 4
```

```python
%who
```

Output:
```
Interactive namespace is empty.
```

```python
var_list = %who_ls
```

```python
var_list
```

Output:
```
['datasets', 'iris']
```

## Part A. Load the Iris dataset.

Sklearn comes with prepared datasets that you can load. To load the iris dataset, run:
```python
from sklearn import datasets
iris = datasets.load_iris()
```

```python
%lsmagic
```

Output:
```
Available line magics:
%alias  %alias_magic  %autoawait  %autocall  %automagic  %autosave  %bookmark  %cd  %clear  %cls  %colors  %conda  %config  %connect_info  %copy  %ddir  %debug  %dhist  %dirs  %doctest_mode  %echo  %ed  %edit  %env  %gui  %hist  %history  %killbgscripts  %ldir  %less  %load  %load_ext  %loadpy  %logoff  %logon  %logstart  %logstate  %logstop  %ls  %lsmagic  %macro  %magic  %matplotlib  %mkdir  %more  %notebook  %page  %pastebin  %pdb  %pdef  %pdoc  %pfile  %pinfo  %pinfo2  %pip  %popd  %pprint  %precision  %prun  %psearch  %psource  %pushd  %pwd  %pycat  %pylab  %qtconsole  %quickref  %recall  %rehashx  %reload_ext  %ren  %rep  %rerun  %reset  %reset_selective  %rmdir  %run  %save  %sc  %set_env  %store  %sx  %system  %tb  %time  %timeit  %unalias  %unload_ext  %who  %who_ls  %whos  %xdel  %xmode

Available cell magics:
%%!  %%HTML  %%SVG  %%bash  %%capture  %%cmd  %%debug  %%file  %%html  %%javascript  %%js  %%latex  %%markdown  %%perl  %%prun  %%pypy  %%python 
```

```python
from sklearn import datasets
iris = datasets.load_iris()
```

```python
import pandas as pd
```

```python
df = pd.DataFrame(iris.data, columns=iris.feature_names)
```

```python
df.head(2)
```

Output:
```
sepal length (cm)  sepal width (cm)  petal length (cm)  petal width (cm)
0                5.1               3.5                1.4               0.2
1                4.9               3.0                1.4               0.2
```

```python
df.columns
```

Output:
```
Index(['sepal length (cm)', 'sepal width (cm)', 'petal length (cm)',
       'petal width (cm)'],
      dtype='object')
```

```python
df['sepal length (cm)'].head()
```

Output:
```
0    5.1
1    4.9
2    4.7
3    4.6
4    5.0
Name: sepal length (cm), dtype: float64
```

```python
df.columns = [s.replace(" ", "_") for s in df.columns]
```

```python
df.columns
```

Output:
```
Index(['sepal_length_(cm)', 'sepal_width_(cm)', 'petal_length_(cm)',
       'petal_width_(cm)'],
      dtype='object')
```

```python
df.columns = [s.replace("(", "").replace(")", "") for s in df.columns]
```

```python
df.columns
```

Output:
```
Index(['sepal_length_cm', 'sepal_width_cm', 'petal_length_cm',
       'petal_width_cm'],
      dtype='object')
```

```python
df.sepal_length_cm.head()
```

Output:
```
0    5.1
1    4.9
2    4.7
3    4.6
4    5.0
Name: sepal_length_cm, dtype: float64
```

```python
from pathlib import Path
```

```python
p = Path()
```

```python
p.absolute()
```

Output:
```
WindowsPath('C:/Users/huy13/Documents/ML/wolfhank/spring2020/01')
```

```python
p.home()
```

Output:
```
WindowsPath('C:/Users/huy13')
```

```python
p2 = p / "02"
```

```python
datafile = p2 / "winequality-red.csv"
```

```python
datafile_path = Path("02/winequality-red.csv")
```

```python
p = %pwd
```

```python
p
```

Output:
```
'C:\\Users\\huy13\\Documents\\ML\\wolfhank\\spring2020\\01'
```

```python
p + "02/" + "winequality-red.csv"
```

```python
%who
```

Output:
```
datasets	 df	 iris	 pd
```

```python
%whos
```

Output:
```
Variable   Type         Data/Info
---------------------------------
datasets   module       <module 'sklearn.datasets<...>\\datasets\\__init__.py'>
df         DataFrame         sepal length (cm)  s<...>n\n[150 rows x 4 columns]
iris       Bunch        {'data': array([[5.1, 3.5<...>atasets\\data\\iris.csv'}
pd         module       <module 'pandas' from 'C:<...>es\\pandas\\__init__.py'>
```

```python
%%javascript
console.log("Hello World")
```

Output:
```
<IPython.core.display.Javascript object>
```

```python
iris.target_names
```

Output:
```
array(['setosa', 'versicolor', 'virginica'], dtype='<U10')
```

```python
iris.target
```

Output:
```
array([0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2])
```

<!-- grade_id: cell-c1e7f79ad332b8f9, points: 0, grade: false, solution: true -->
```python
from sklearn import datasets

# Load the iris dataset
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-fd50146eeead0c65, points: 1, grade: true, solution: false -->
```python
assert (iris['data'] == datasets.load_iris()['data']).all()
```

## Part B. Use an SVM to predict iris species.

To setup the data and train an SVM, run the following:
```python
from sklearn.svm import SVC

X = iris.data
y = iris.target

classifier = SVC(kernel="rbf")
classifier.fit(X, y)
predictions = classifier.predict(X)
```

```python
%who
```

Output:
```
SVC	 X	 classifier	 datasets	 iris	 predictions	 var_list	 y
```

pd.DataFrame(iris.data, columns = iris.feature_names)

```python
from sklearn.svm import SVC
```

```python
svc = SVC()
```

```python
svc
```

Output:
```
SVC(C=1.0, cache_size=200, class_weight=None, coef0=0.0,
    decision_function_shape='ovr', degree=3, gamma='auto_deprecated',
    kernel='rbf', max_iter=-1, probability=False, random_state=None,
    shrinking=True, tol=0.001, verbose=False)
```

```python
X = iris.data
y = iris.target
```

X = iris.data
y = pd.Series(iris.target)

```python
type(iris.target)
```

Output:
```
numpy.ndarray
```

```python
iris.target.shape
```

Output:
```
(150,)
```

```python
df.shape
```

Output:
```
(150, 4)
```

```python
classifier = SVC()
```

```python
X.shape

classifier.fit(X, y)
```

Output:
```
C:\Users\huy13\AppData\Local\Continuum\anaconda3\lib\site-packages\sklearn\svm\base.py:193: FutureWarning: The default value of gamma will change from 'auto' to 'scale' in version 0.22 to account better for unscaled features. Set gamma explicitly to 'auto' or 'scale' to avoid this warning.
  "avoid this warning.", FutureWarning)
SVC(C=1.0, cache_size=200, class_weight=None, coef0=0.0,
    decision_function_shape='ovr', degree=3, gamma='auto_deprecated',
    kernel='rbf', max_iter=-1, probability=False, random_state=None,
    shrinking=True, tol=0.001, verbose=False)
```

```python
classifier.score(X, y)
```

Output:
```
0.9866666666666667
```

```python
predictions = classifier.predict(X)
```

- Overfitting (Variance): tailored to a specific dataset (does not generalize well), misses the point, memorizes the data
- Underfitting (Bias): accuracy could be better, too general 

1. Import
2. Instantiate (a class)
  - Estimators:
    - Classifiers
    - Regressors
3. Fit
4. Predict

```python
predictions
```

Output:
```
array([0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 2, 1, 1, 1, 1, 1, 2, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2])
```

```python
from sklearn.model_selection import train_test_split
```

<!-- grade_id: cell-bc8ce57a60e5fd7a, points: 0, grade: false, solution: true -->
```python
from sklearn.svm import SVC

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-165c467d2a6e7dde, points: 1, grade: true, solution: false -->
```python
assert X.shape == (150, 4)
assert len(y) == 150
```